```python
# Initialize Otter
import otter
grader = otter.Notebook("h01-python.ipynb")
```

# h01. Python Lists, Looping, and Function

### A Quick Note on How Grading Works

Please read this carefully to understand the tests in this notebook.

* **Public Tests:** The instant checks you can run here are "public tests." They are designed to help you make sure your answer is in the correct format (for example, you've named your variable correctly). Passing them is a great first step!

* **Hidden Tests:** Your final grade is determined by a more complete set of "hidden tests." They are more thorough and check if your code is truly correct.

**Bottom line:** Passing the public tests does **not** guarantee a perfect score. Always be sure to double-check your logic and test your code yourself.

```python
pip install otter-grader
```

**Question 1: Creating a List of Fruits**

Create a new Python list named `fruits`. The list must contain the following 5 string elements in this exact order:

1.  `"apple"`
2.  `"banana"`
3.  `"orange"`
4.  `"grape"`
5.  `"strawberry"`

```python
fruits = ...
fruits
```

```python
grader.check("Q01")
```

### Q02. Access and print the first and last elements of your fruit list.

```python
first_fruit = ...
last_fruit = ...
first_fruit
last_fruit
```

```python
grader.check("Q02")
```

### Q03. Add a new fruit mango to the end of the list 
- using the append() method.


```python
...
```

```python
grader.check("Q03")
```

### Q04. Insert a fruit durian at the second position in the list.

```python
...
```

```python
grader.check("Q04")
```

### Q05. Remove the fruit orange by its name 
- using the remove() method.

```python
...
```

```python
grader.check("Q05")
```

### Q06. Write a loop to iterate over the list and print each fruit in upper case.
- using the upper() method 

```python
...
    ...
```

### Q07. Write a for loop to count the number of the elements in the list.
- use a "counter" variable
- create a loop variable
- use the in keyword before "fruits"
- increment the counter inside the loop

```python
counter = ...
...
    ...
```

```python
grader.check("Q07")
```

### Q08. Check if 'apple' exists in the list and print a message "apple found" if it does.

```python
...
    ...
```

**Question 9: Create Sorted Versions of a List**

Using the original `fruits` list, perform the following two steps:

1.  Create a **new list** called `sorted_list` that contains all the fruits sorted alphabetically (A-Z).
2.  Create a **second new list** called `reverse_sorted_list` that contains all the fruits sorted in reverse-alphabetical order (Z-A).



```python
sorted_list = ...
reverse_sorted_list = ...
sorted_list
reverse_sorted_list
```

```python
grader.check("Q09")
```

### Q10. Design a function:
- call it count_elements
- with a parameter call lst (which means list)
- when called with a list argument (this fruits list), count the number of elements in the list
- return the number of elements in the list to the caller

```python
def count_elements(lst):
    number_of_elements = ...
    ...

fruit_count= ...
```

```python
grader.check("Q10")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for Q01 (points: [0, 0])
```
>>> assert isinstance(fruits, list), "The variable 'fruits' should be a list."
```
```
>>> assert len(fruits) == 5, "The 'fruits' list should contain exactly 5 items."
```


Tests for Q02 (points: [0, 0, 0])
```
>>> assert 'fruits' in globals(), "This test requires the 'fruits' list from Q01."
```
```
>>> assert first_fruit == fruits[0], f"Expected first_fruit to be '{fruits[0]}', but got '{first_fruit}'."
```
```
>>> assert last_fruit == fruits[-1], f"Expected last_fruit to be '{fruits[-1]}', but got '{last_fruit}'."
```


Tests for Q03 (points: [0, 0])
```
>>> assert len(fruits) == 6, "After appending, the 'fruits' list should have 6 items."
```
```
>>> assert fruits[-1] == 'mango', "The last item in the list should now be 'mango'."
```


Tests for Q04 (points: [0, 0])
```
>>> assert len(fruits) == 7, "After inserting an element, the 'fruits' list should have 7 items."
```
```
>>> assert fruits[1] == 'durian', "The item at index 1 should now be 'durian'."
```


Tests for Q05 (points: [0, 0, 0])
```
>>> assert len(fruits) == 6, "After removing an element, the 'fruits' list should have 6 items."
```
```
>>> assert 'orange' not in fruits, "The string 'orange' should no longer be in the fruits list."
```
```
>>> assert fruits[3] == 'grape', "The element that came after 'orange' should now be at index 3."
```


Tests for Q07 (points: [0, 0])
```
>>> assert counter == 6, f"Expected the counter to be 6, but got {counter} instead."
```
```
>>> assert counter == len(fruits), "The final value of 'counter' should be equal to the number of elements in the 'fruits' list."
```


Tests for Q09 (points: [0, 0])
```
>>> expected_sorted = ['apple', 'banana', 'durian', 'grape', 'mango', 'strawberry']
>>> assert sorted_list == expected_sorted, "The 'sorted_list' is not sorted correctly in alphabetical order (A-Z)."
```
```
>>> expected_reverse_sorted = ['strawberry', 'mango', 'grape', 'durian', 'banana', 'apple']
>>> assert reverse_sorted_list == expected_reverse_sorted, "The 'reverse_sorted_list' is not sorted correctly in reverse-alphabetical order (Z-A)."
```


Tests for Q10 (points: [0, 0, 0])
```
>>> from types import FunctionType
>>> result_general = count_elements([1, 2, 3, 4, 5, 6, 7, 8])
>>> assert result_general == 8, "The function did not return the correct count for a list of 8 numbers."
```
```
>>> result_empty = count_elements([])
>>> assert result_empty == 0, "The function should return 0 when called with an empty list."
```
```
>>> assert 'fruit_count' in globals(), "The variable 'fruit_count' should be defined."
>>> 
>>> assert fruit_count == 6, f"Expected fruit_count to be 6, but got {fruit_count} instead."
>>> assert isinstance(fruit_count, int), "The 'fruit_count' variable should be an integer."
>>>
```
