```python
# Initialize Otter
import otter
grader = otter.Notebook("h03-numpy.ipynb")
```

# h03 NumPy
---

### List Operations

### NumPy Operations

```python
### import numpy as np
import numpy as np
```

### 1. Review list operations

### Task: create a list of elements 0 to 9:
1.  Create an empty list variable, call it "lst".
2.  Use a `for` loop.
3.  Use the `range` function.
4.  Use the `append()` method.

```python
lst = []
...
    ...
```

```python
### evaluate the list
lst
```

```python
grader.check("Q01_1")
```

**check the data type of the list**


```python
type(lst)
```

```python
grader.check("Q01_2")
```

**append() 6 to list "lst"**


```python
lst_append = ...
...
lst
```

```python
grader.check("Q01_3")
```

**extend lst using extend() to add 11, 12, 13 to lst**`


```python
...
lst
```

```python
grader.check("Q01_4")
```

### write a for loop to remove some elements
1.  Use the `range` function.
2.  Use the `pop()` method to remove 6, 7, 8, 9 from the list.
3.  *note: pop() removes the last element by default if you don't assign an index.*

```python
lst = ...
lst
```

```python
grader.check("Q01_5")
```

### write a for loop to remove some elements
1.  Use the `range` function.
2.  Use the `pop()` method to remove 6, 7, 8, 9 from the list.
3.  *note: pop() removes the last element by default if you don't assign an index.*

```python
for i in range(...):
    lst....

lst
```

```python
grader.check("Q01_6")
```

### 2. Create a numpy array

```python
### Use the numpy array function (if it's a method, you have to use it with an object)
### turn the list lst into a numpy array
### assign the numpy array to a variable, call it "arr"

arr = np.array(lst)
```

```python
### evaluate evaluate the numpy array arr, 

arr
```

```python
### Now we have a list (lst) and a numpy array (arr). 
### They have the same literal value elements, 
### but they are of different data types
```

```python
### check the data type of arr
### you should see the data type is "numpy.ndarray"

type(arr)
```

```python
lst
```

```python

```

```python

```

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for Q01_1 (points: [0, 0])
```
>>> assert len(lst) == 10, f'The list should have 10 elements, but yours has {len(lst)}.'
```
```
>>> expected_list = [0, 1, 2, 3, 4, 5, 6, 7, 8, 9]
>>> assert lst == expected_list, f'The list should be {expected_list}, but yours is {lst}.'
```


Tests for Q01_2 (points: [0])
```
>>> assert isinstance(lst, list), "'lst' should be a list, but it is a " + str(type(lst))
```


Tests for Q01_3 (points: [0, 0, 0])
```
>>> assert len(lst) == 11, f'The list should have 11 elements after appending, but yours has {len(lst)}.'
```
```
>>> assert lst[-1] == 6, f'The last element in the list should be 6, but yours is {lst[-1]}.'
```
```
>>> expected_list = [0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 6]
>>> assert lst == expected_list, "The final list's contents are incorrect."
```


Tests for Q01_4 (points: [0, 0])
```
>>> assert len(lst) == 14, f'The list should have 13 elements, but yours has {len(lst)}.'
```
```
>>> expected_list = [0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 6, 11, 12, 13]
>>> assert lst == expected_list, "The contents of 'lst' are incorrect after extending."
```


Tests for Q01_5 (points: [0, 0])
```
>>> assert len(lst) == 10, "The length of 'lst' is incorrect after slicing. It should have 10 elements."
```
```
>>> expected_list = [0, 1, 2, 3, 4, 5, 6, 7, 8, 9]
>>> assert lst == expected_list, "The elements in 'lst' are incorrect after slicing."
```


Tests for Q01_6 (points: [0, 0])
```
>>> assert len(lst) == 6, "The length of 'lst' is incorrect. It should have 6 elements after removing four."
```
```
>>> expected_list = [0, 1, 2, 3, 4, 5]
>>> assert lst == expected_list, "The final contents of 'lst' are incorrect. The elements 6, 7, 8, and 9 should have been removed."
```
