```python
# Initialize Otter
import otter
grader = otter.Notebook("hw03-numpy.ipynb")
```

# h03 NumPy - Questions
In this homework, you will do:
- Python review: for loop & data collection operations
- NumPy array operations
- NumPy vectorized functions (**ufuncs**)

**Academic honor**:
1. Academic honor is important. Do your work, learn, **DON'T CHEAT!** 
2. <font color='darkred'>**Use AI to LEARN, not to cheat**. **DO NOT do copy-and-paste**</font>. 
3. Collaboration and using resources are encouraged, but DO NOT directly share answers.

# Notebook Editing Rules

1. **Only modify code where you see**  
   `##### change me #####` *or a similar placeholder.*  
   Write your solution on that line.

2. **Do NOT rename any variables already provided** in the prompts or tests  
   (e.g., `arr1_mean`, `subset`, `row1`).  
   The autograder looks for those exact names.

Run each cell after editing to confirm it executes without errors.


```python
%pip install numpy
```

```python
%pip show numpy
```

```python
import numpy as np
```

## 1. Programming Review
### 1.1. for loop with range

```python
### use the range( ) function to print numbers from 0 to 9
### use the end= parameter when printing

##### change me #####:
    print(i, end=" ")
```

### 1.2. for loop with list

```python
### use the range() function to create the sequence from 0 to 9
### use the list function to cast the object list
### save the list to a variable named "lst"
### use a for loop to print the list elements (use the "in" keyword)
### use the end= parameter when printing

lst = ##### change me #####
for ##### change me #####:
    print(i, end=" ")
```

## 1.3 dictionary

```python
### Observe the hypothetical pocket change dictionary below. Note that a dictionary:
### use curly braces instead of square brackets like lists
### contains key: value pairs as elements

dict_money = {"Garrett": 1900, "Jase": 20000, "Quincy": 23000, "Canaan": 21000, "Rhea": 19000}
```

To access the value of a key, use the **bracket notation**.

```python
dict_money['Garrett']
```

```python
total_money= ##### change me #####
total_money
```

```python
grader.check("Q01_3")
```

## 2. NumPy Array Operations

```python
### import numpy here if you haven't
import numpy as np
```

### 2.1 Creating NumPy Arrays: 1D (vector) & 2D (matrix)
Practicing array creation and attributes: 
- **dtype**,
- **ndim**, and
- **shape**.

```python
### Use the numpy arrange function to to build a one-dimensional array containing the integers 0 through 9 in order.
### assign the numpy array to a variable, call it "arr"

arr = ##### change me #####
arr
```

```python
### check the dimension of the array and store it to a variable dimension
### use the ndmin attribute

dimension= arr.##### change me #####
dimension
```

```python
grader.check("Q02_1")
```

```python
### Create a 3 × 3 NumPy array that looks exactly like this:
[[1, 2, 3],
 [5, 5, 5],
 [7, 8, 9]]

### Store the array in a variable named arr_2d.

arr_2d = np.array([[1, 2, 3],
                ##### change me #####
                ##### change me #####
arr_2d
```

```python
### check the shape of the array
### use the shape attribute

##### change me #####
```

```python
grader.check("Q02_2")
```

```python
##Use `type()` to find out what Python thinks the object **`arr_2d`** is.  
###Store the result in the variable **`arr_type`**.  
### you should see the data type is "numpy.ndarray"

arr_type = ##### change me #####(arr_2d)
arr_type
```

```python
###Use the `.dtype` attribute to discover the data type of the *elements* inside `arr_2d`.  
###Store that result in **`scalar_dtype`**.  
### you should see dtype('int64')
scalar_type = arr_2d.##### change me #####
scalar_type
```

```python
grader.check("Q02_3")
```

```python
### create a 1D array with 24 elements 
### the elements are random float numbers between [0, 1). 
### use one of the random functions in the np.random module

arr = np.##### change me #####
arr
```

```python
grader.check("Q02_4")
```

```python
### reshape the array into a 6x4 2D array

arr_64 = arr.##### change me #####
arr_64
```

```python
grader.check("Q02_5")
```

```python
### create a 1D array with 24 elements 
### the elements are random integers between [1, 100] (including 100)
### use one of the random functions in the np.random module


np.random.seed(100)     ### don't change the seed value
                        ### with the same seed value, the "randomly" 
                        ### generated values will be the same every time you run the code

arr_int = np.random.##### change me #####
arr_int
```

```python
grader.check("Q02_6")
```

```python
### reshape the 1D array into a 6x4 2D array, still call it arr_int

arr_int = arr_int.##### change me #####
arr_int
```

```python
grader.check("Q02_7")
```

```python
### show the first row of the array arr_int

arr_int[0]
```

```python
### Slicing task – the two middle rows

###`arr_int` is a 6 × 4 array from the question Q02_7.  
###Slice **rows 2 and 3** (0-indexed) and store that 2 × 4 result in a variable named **`mid_rows`**.

mid_rows= arr_int[ ##### change me ##### ]  
mid_rows
```

```python
grader.check("Q02_8")
```

```python
### Slicing task – select a subset

###Slice the following subset from `arr_int` and store it in a variable called subset.  

# array([[99, 54, 67],
#        [15, 35, 25],
#        [61, 59, 17],
#        [94, 87,  3]])


subset = arr_int[ ##### change me ##### ]
subset
```

```python
grader.check("Q02_9")
```

```python
arr1 = np.arange(5)           ### 0...4
arr2 = np.arange(5, 10, 1)    ### 5...9
print("arr1 = ", arr1)
print("arr2 = ", arr2)
```

```python
### addition: add arr1 and arr2 and store it in a variable sum_arr

sum_arr = ##### change me #####
sum_arr
```

```python
grader.check("Q03_1")
```

```python
### subtraction: subtract arr2 from arr1 and store the result in a variable difference

difference= ##### change me #####
difference
```

```python
grader.check("Q03_2")
```

```python
### multiplication: times arr1 by arr2 and store it in a variable multiplied

multiplied= ##### change me #####
multiplied
```

```python
grader.check("Q03_3")
```

```python
### division: divide arr1 by arr2 and store it in a variable divided

divided = ##### change me #####
divided
```

```python
grader.check("Q03_4")
```

```python
### min: check the minimum value of arr1 and store it in a variable min_value

min_value = ##### change me #####
min_value
```

```python
grader.check("Q03_5")
```

```python
### max: check the maximum value of arr1 and store it in a variable max_value

max_value = ##### change me #####
max_value
```

```python
grader.check("Q03_6")
```

```python
### mean: check the average of elements in arr1

arr1_mean = ##### change me #####
arr1_mean
```

```python
grader.check("Q03_7")
```

```python
### argmax(): check the index number of the largest value in arr1

index = ##### change me #####
index
```

```python
grader.check("Q03_8")
```

```python
### sqrt(): take square root of arr1

square_root = ##### change me #####
square_root
```

```python
grader.check("Q03_9")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for Q01_3 (points: [0, 0])
```
>>> assert isinstance(total_money, int), "total_money must be an integer."
```
```
>>> expected = dict_money["Canaan"] + dict_money["Rhea"]
>>> assert total_money == expected, f"total_money is {total_money}, but it should be {expected}."
```


Tests for Q02_1 (points: [0, 0, 0, 0])
```
>>> import numpy as np
>>> assert isinstance(arr, np.ndarray), "arr must be a NumPy ndarray."
>>>
```
```
>>> assert arr.size == 10, f"arr should contain 10 elements, found {arr.size}."
```
```
>>> expected = np.arange(10)
>>> assert np.array_equal(arr, expected),  f"arr is {arr}, but it should be {expected}."
>>>
```
```
>>> assert dimension == arr.ndim == 1, f"dimension is {dimension}, but arr.ndim reports {arr.ndim} (expected 1)."
>>>
```


Tests for Q02_2 (points: [0, 0, 0])
```
>>> import numpy as np
>>> assert isinstance(arr_2d, np.ndarray), "arr_2d must be a NumPy array."
>>>
```
```
>>> assert arr_2d.shape == (3, 3), f"arr_2d.shape is {arr_2d.shape}, expected (3, 3)."
```
```
>>> expected = np.array([[1, 2, 3],
...                      [5, 5, 5],
...                      [7, 8, 9]])
>>> assert np.array_equal(arr_2d, expected), f"arr_2d is\n{arr_2d}\n\nbut should be\n{expected}"
```


Tests for Q02_3 (points: [0, 0, 0, 0])
```
>>> import numpy as np
>>> assert arr_type is type(arr_2d), f"arr_type is {arr_type}, but type(arr_2d) is {type(arr_2d)}."
```
```
>>> assert arr_type.__name__ == "ndarray", "arr_type should be numpy.ndarray."
```
```
>>> assert isinstance(scalar_type, np.dtype), "scalar_type must be a NumPy dtype."
>>>
```
```
>>> assert scalar_type == arr_2d.dtype, f"scalar_type is {scalar_type}, but arr_2d.dtype is {arr_2d.dtype}."
>>>
```


Tests for Q02_4 (points: [0, 0, 0, 0, 0])
```
>>> import numpy as np
>>> assert isinstance(arr, np.ndarray), "arr must be a NumPy array."
>>>
```
```
>>> assert arr.size == 24, f"arr should have 24 elements, found {arr.size}."
>>>
```
```
>>> assert arr.ndim == 1,  f"arr.ndim is {arr.ndim}, expected 1."
```
```
>>> assert arr.min() >= 0, f"Minimum value {arr.min()} is below 0."
>>>
```
```
>>> assert arr.max() <  1, f"Maximum value {arr.max()} is not less than 1."
>>>
```


Tests for Q02_5 (points: [0, 0])
```
>>> import numpy as np
>>> assert arr_64.shape == (6, 4), f"arr_64.shape is {arr_64.shape}, expected (6, 4)."
>>>
```
```
>>> assert np.array_equal(arr_64.flatten(), arr), "The elements in arr_64 do not match the values in the original arr."
```


Tests for Q02_6 (points: [0, 0, 0])
```
>>> import numpy as np
>>> assert arr_int.size == 24, f"arr_int should contain 24 elements, found {arr_int.size}."
>>>
```
```
>>> assert arr_int.min() >= 1,  f"Minimum value {arr_int.min()} is below 1."
>>>
```
```
>>> assert arr_int.max() <= 100, f"Maximum value {arr_int.max()} is above 100."
>>>
```


Tests for Q02_7 (points: [0, 0])
```
>>> import numpy as np
>>> assert arr_int.size == 24, "arr_int should contain 24 elements after reshaping."
```
```
>>> assert arr_int.shape == (6, 4), f"arr_int.shape is {arr_int.shape}, expected (6, 4)."
```


Tests for Q02_8 (points: [0, 0])
```
>>> import numpy as np
>>> expected = arr_int[2:4, :]
>>> assert np.array_equal(mid_rows, expected), f"mid_rows is\n{mid_rows}\nexpected\n{expected}"
```
```
>>> assert mid_rows.shape == (2, 4), f"mid_rows.shape is {mid_rows.shape}, expected (2, 4)."
```


Tests for Q02_9 (points: [0, 0])
```
>>> import numpy as np
>>> expected = np.array([
...     [99, 54, 67],
...     [15, 35, 25],
...     [61, 59, 17],
...     [94, 87,  3]])
>>> 
>>> assert np.array_equal(subset, expected), f"subset is\n{subset}\nexpected\n{expected}"
```
```
>>> assert subset.shape == (4, 3), f"subset.shape is {subset.shape}, expected (4, 3)."
```


Tests for Q03_1 (points: [0])
```
>>> import numpy as np
>>> expected = arr1 + arr2  
>>> assert np.array_equal(sum_arr, expected), f"sum_arr is {sum_arr}, expected {expected}."
>>>
```


Tests for Q03_2 (points: [0])
```
>>> import numpy as np
>>> expected = arr1 - arr2
>>> assert np.array_equal(difference, expected), f"difference is {difference}, expected {expected}."
>>>
```


Tests for Q03_3 (points: [0])
```
>>> import numpy as np
>>> 
>>> expected = arr1 * arr2
>>> assert np.array_equal(multiplied, expected), f"multiplied is {multiplied}, expected {expected}."
>>>
```


Tests for Q03_4 (points: [0])
```
>>> import numpy as np
>>> exp_divided = arr1 / arr2
>>> assert np.allclose(divided, exp_divided),  f"divided is {divided}, expected {exp_divided}."
```


Tests for Q03_5 (points: [0])
```
>>> import numpy as np
>>> exp_min = arr1.min()  
>>> assert min_value == exp_min, f"min_value is {min_value}, expected {exp_min}."
>>>
```


Tests for Q03_6 (points: [0])
```
>>> import numpy as np
>>> expected_max = arr1.max()          
>>> assert max_value == expected_max, f"max_value is {max_value}, expected {expected_max}."
>>>
```


Tests for Q03_7 (points: [0])
```
>>> import numpy as np
>>> expected_mean = arr1.mean()
>>> assert np.isclose(arr1_mean, expected_mean), f"arr1_mean is {arr1_mean}, expected {expected_mean}."
>>>
```


Tests for Q03_8 (points: [0])
```
>>> import numpy as np
>>> expected_idx = arr1.argmax()
>>> assert index == expected_idx, f"index is {index}, expected {expected_idx}."
```


Tests for Q03_9 (points: [0])
```
>>> import numpy as np
>>> expected_sqrt = np.sqrt(arr1)
>>> assert np.allclose(square_root, expected_sqrt), f"square_root is {square_root}, expected {expected_sqrt}."
>>>
```
